Before you submit this notebook, make sure everything runs as expected in the local test cases. 
Please, paste the solution to the designed cell and do not change anything else.

Also, please, leave your first and last names below

```python
FirstName = ""
LastName = ""
```

---

```python
import numpy as np
```

<!-- grade_id: cell-ac8fc52d8a39ccb4, points: 0, grade: false, solution: true -->
```python
class Module(object):
    """
    Basically, you can think of a module as of a something (black box) 
    which can process `input` data and produce `ouput` data.
    This is like applying a function which is called `forward`: 
        
        output = module.forward(input)
    
    The module should be able to perform a backward pass: to differentiate the `forward` function. 
    More, it should be able to differentiate it if is a part of chain (chain rule).
    The latter implies there is a gradient from previous step of a chain rule. 
    
        gradInput = module.backward(input, gradOutput)
    """
    def __init__ (self):
        self.output = None
        self.gradInput = None
        self.training = True
    
    def forward(self, input):
        """
        Takes an input object, and computes the corresponding output of the module.
        """
        return self.updateOutput(input)

    def backward(self,input, gradOutput):
        """
        Performs a backpropagation step through the module, with respect to the given input.
        
        This includes 
         - computing a gradient w.r.t. `input` (is needed for further backprop),
         - computing a gradient w.r.t. parameters (to update parameters while optimizing).
        """
        self.updateGradInput(input, gradOutput)
        self.accGradParameters(input, gradOutput)
        return self.gradInput
    

    def updateOutput(self, input):
        """
        Computes the output using the current parameter set of the class and input.
        This function returns the result which is stored in the `output` field.
        
        Make sure to both store the data in `output` field and return it. 
        """
        
        # The easiest case:
            
        # self.output = input 
        # return self.output
        
        pass

    def updateGradInput(self, input, gradOutput):
        """
        Computing the gradient of the module with respect to its own input. 
        This is returned in `gradInput`. Also, the `gradInput` state variable is updated accordingly.
        
        The shape of `gradInput` is always the same as the shape of `input`.
        
        Make sure to both store the gradients in `gradInput` field and return it.
        """
        
        # The easiest case:
        
        # self.gradInput = gradOutput 
        # return self.gradInput
        
        pass   
    
    def accGradParameters(self, input, gradOutput):
        """
        Computing the gradient of the module with respect to its own parameters.
        No need to override if module has no parameters (e.g. ReLU).
        """
        pass
    
    def zeroGradParameters(self): 
        """
        Zeroes `gradParams` variable if the module has params.
        """
        pass
        
    def getParameters(self):
        """
        Returns a list with its parameters. 
        If the module does not have parameters return empty list. 
        """
        return []
        
    def getGradParameters(self):
        """
        Returns a list with gradients with respect to its parameters. 
        If the module does not have parameters return empty list. 
        """
        return []
    
    def train(self):
        """
        Sets training mode for the module.
        Training and testing behaviour differs for Dropout, BatchNorm.
        """
        self.training = True
    
    def evaluate(self):
        """
        Sets evaluation mode for the module.
        Training and testing behaviour differs for Dropout, BatchNorm.
        """
        self.training = False
    
    def __repr__(self):
        """
        Pretty printing. Should be overrided in every module if you want 
        to have readable description. 
        """
        return "Module"
    
class Sequential(Module):
    """
         This class implements a container, which processes `input` data sequentially. 
         
         `input` is processed by each module (layer) in self.modules consecutively.
         The resulting array is called `output`. 
    """
    
    def __init__ (self):
        super(Sequential, self).__init__()
        self.modules = []
   
    def add(self, module):
        """
        Adds a module to the container.
        """
        self.modules.append(module)

    def updateOutput(self, input):
        """
        Basic workflow of FORWARD PASS:
        
            y_0    = module[0].forward(input)
            y_1    = module[1].forward(y_0)
            ...
            output = module[n-1].forward(y_{n-2})   
            
            
        Just write a little loop. 
        """

        # Your code goes here. ################################################
        self.input_lst = []
        current_input = input
        
        for module in self.modules:
            current_output = module.forward(current_input)
            self.input_lst.append(current_input)
            current_input = current_output

        self.output = current_output
        self.input_lst = self.input_lst[1:]
        
        # print(current_output)
        # print(self.input_lst)

        return self.output

    def backward(self, input, gradient_output):
        """
        Workflow of BACKWARD PASS:
            
            g_{n-1} = module[n-1].backward(y_{n-2}, gradOutput)
            g_{n-2} = module[n-2].backward(y_{n-3}, g_{n-1})
            ...
            g_1 = module[1].backward(y_0, g_2)   
            gradInput = module[0].backward(input, g_1)   
             
             
        !!!
                
        To ech module you need to provide the input, module saw while forward pass, 
        it is used while computing gradients. 
        Make sure that the input for `i-th` layer the output of `module[i]` (just the same input as in forward pass) 
        and NOT `input` to this Sequential module. 
        
        !!!
        
        """
        # Your code goes here. ################################################
        
        current_inputs = list(reversed(self.input_lst)) + [input]
        current_modules = list(reversed(self.modules))
        
        for i in range(len(current_inputs)):
            current_input = current_inputs[i]
            current_module = current_modules[i]
            gradient_input = current_module.backward(current_input, gradient_output)
            gradient_output = gradient_input
        self.gradInput = gradient_input
        
        return self.gradInput
      

    def zeroGradParameters(self): 
        for module in self.modules:
            module.zeroGradParameters()
    
    def getParameters(self):
        """
        Should gather all parameters in a list.
        """
        return [x.getParameters() for x in self.modules]
    
    def getGradParameters(self):
        """
        Should gather all gradients w.r.t parameters in a list.
        """
        return [x.getGradParameters() for x in self.modules]
    
    def __repr__(self):
        string = "".join([str(x) + '\n' for x in self.modules])
        return string
    
    def __getitem__(self,x):
        return self.modules.__getitem__(x)
    
    def train(self):
        """
        Propagates training parameter through all modules
        """
        self.training = True
        for module in self.modules:
            module.train()
    
    def evaluate(self):
        """
        Propagates training parameter through all modules
        """
        self.training = False
        for module in self.modules:
            module.evaluate()
            
            
class Linear(Module):
    """
    A module which applies a linear transformation 
    A common name is fully-connected layer, InnerProductLayer in caffe. 
    
    The module should work with 2D input of shape (n_samples, n_feature).
    """
    def __init__(self, n_in, n_out):
        super(Linear, self).__init__()
       
        # This is a nice initialization
        stdv = 1./np.sqrt(n_in)
        self.W = np.random.uniform(-stdv, stdv, size = (n_out, n_in))
        self.b = np.random.uniform(-stdv, stdv, size = n_out)
        
        self.gradW = np.zeros_like(self.W)
        self.gradb = np.zeros_like(self.b)
        
    def updateOutput(self, input):
        # Your code goes here. ################################################
        self.output = np.dot(input, self.W.T) + self.b

        return self.output
    
    def updateGradInput(self, input, gradOutput):
        # Your code goes here. ################################################
        self.gradInput = np.dot(gradOutput, self.W)

        return self.gradInput
    
    def accGradParameters(self, input, gradOutput):
        # Your code goes here. ################################################
        self.gradW = np.dot(gradOutput.T, input)
        self.gradb = np.sum(gradOutput, axis=0)
    
    def zeroGradParameters(self):
        self.gradW.fill(0)
        self.gradb.fill(0)
        
    def getParameters(self):
        return [self.W, self.b]
    
    def getGradParameters(self):
        return [self.gradW, self.gradb]
    
    def __repr__(self):
        s = self.W.shape
        q = 'Linear %d -> %d' %(s[1],s[0])
        return q
    
    
class SoftMax(Module):
    def __init__(self):
         super(SoftMax, self).__init__()
    
    def updateOutput(self, input):
        # start with normalization for numerical stability
        self.output = np.subtract(input, input.max(axis=1, keepdims=True))
        
        # Your code goes here. ################################################
        self.output = np.exp(self.output) / np.sum(np.exp(self.output), axis=1, keepdims=True)
        
        return self.output
    
    def updateGradInput(self, input, gradOutput):
        # Your code goes here. ################################################
        pred_grad = np.multiply(self.output, gradOutput)
        self.gradInput = pred_grad - np.multiply(self.output, np.sum(pred_grad, axis=1, keepdims=True))

        return self.gradInput
    
    def __repr__(self):
        return "SoftMax"
    

class LogSoftMax(Module):
    def __init__(self):
         super(LogSoftMax, self).__init__()
    
    def updateOutput(self, input):
        # start with normalization for numerical stability
        self.output = np.subtract(input, input.max(axis=1, keepdims=True))
        
        # Your code goes here. ################################################
        self.output = self.output - np.log(np.sum(np.exp(self.output), axis=1, keepdims=True))
        # print(self.output.shape)
        
        return self.output
    
    def updateGradInput(self, input, gradOutput):
        # Your code goes here. ################################################
        current_input = np.exp(input - input.max(axis=1, keepdims=True))
        # print(current_input.shape)
        current_input /= np.sum(current_input, axis=1, keepdims=True)
        self.gradInput = gradOutput - np.multiply(current_input, np.sum(gradOutput, axis=1, keepdims=True))
        return self.gradInput
    
    def __repr__(self):
        return "LogSoftMax"
    
    
class BatchNormalization(Module):
    EPS = 1e-3
    def __init__(self, alpha = 0.):
        super(BatchNormalization, self).__init__()
        self.alpha = alpha
        self.moving_mean = None 
        self.moving_variance = None
        
    def updateOutput(self, input):
        # Your code goes here. ################################################
        # use self.EPS please
        
        if (self.training == True):
            self.mean = np.mean(input, axis=0, keepdims=True)
            self.var = np.var(input, axis=0, keepdims=True)
            self.moving_mean = self.moving_mean * self.alpha + self.mean * (1 - self.alpha)
            self.moving_variance = self.moving_variance * self.alpha + self.var * (1 - self.alpha)
            self.output = (input - self.mean)/np.sqrt(self.var + self.EPS)
        else:
            self.output = (input - self.moving_mean) / np.sqrt(self.moving_variance + self.EPS)

        return self.output
    
    def updateGradInput(self, input, gradOutput):
        # Your code goes here. ################################################
        
        self.gradInput = input.shape[0] * gradOutput
        self.gradInput -= np.sum(gradOutput, axis=0, keepdims=True)
        self.gradInput -= self.output * np.sum(gradOutput * self.output, axis=0, keepdims=True)
        self.gradInput = 1/(input.shape[0] * np.sqrt(self.var + self.EPS)) * self.gradInput
        
        return self.gradInput
    
    def __repr__(self):
        return "BatchNormalization"
    

class ChannelwiseScaling(Module):
    """
       Implements linear transform of input y = \gamma * x + \beta
       where \gamma, \beta - learnable vectors of length x.shape[-1]
    """
    def __init__(self, n_out):
        super(ChannelwiseScaling, self).__init__()

        stdv = 1./np.sqrt(n_out)
        self.gamma = np.random.uniform(-stdv, stdv, size=n_out)
        self.beta = np.random.uniform(-stdv, stdv, size=n_out)
        
        self.gradGamma = np.zeros_like(self.gamma)
        self.gradBeta = np.zeros_like(self.beta)

    def updateOutput(self, input):
        self.output = input * self.gamma + self.beta
        return self.output
        
    def updateGradInput(self, input, gradOutput):
        self.gradInput = gradOutput * self.gamma
        return self.gradInput
    
    def accGradParameters(self, input, gradOutput):
        self.gradBeta = np.sum(gradOutput, axis=0)
        self.gradGamma = np.sum(gradOutput*input, axis=0)
    
    def zeroGradParameters(self):
        self.gradGamma.fill(0)
        self.gradBeta.fill(0)
        
    def getParameters(self):
        return [self.gamma, self.beta]
    
    def getGradParameters(self):
        return [self.gradGamma, self.gradBeta]
    
    def __repr__(self):
        return "ChannelwiseScaling"
    
    
class Dropout(Module):
    def __init__(self, p=0.5):
        super(Dropout, self).__init__()
        
        self.p = p
        self.mask = None
        
    def updateOutput(self, input):
        # Your code goes here. ################################################
        
        if (self.training == True):
            if (self.p == 0.0):
                self.mask = np.ones_like(input)
            else:
                self.mask = np.random.binomial(1, 1 - self.p, input.shape)
            self.output = np.multiply(input, self.mask) / (1 - self.p)
        else:
            self.output = input
        
        return  self.output
    
    def updateGradInput(self, input, gradOutput):
        # Your code goes here. ################################################
        
        self.gradInput = gradOutput * self.mask / (1 - self.p)
        return self.gradInput
        
        return self.gradInput
        
    def __repr__(self):
        return "Dropout"
    
    
class ReLU(Module):
    def __init__(self):
         super(ReLU, self).__init__()
    
    def updateOutput(self, input):
        self.output = np.maximum(input, 0)
        return self.output
    
    def updateGradInput(self, input, gradOutput):
        self.gradInput = np.multiply(gradOutput , input > 0)
        return self.gradInput
    
    def __repr__(self):
        return "ReLU"
    
class LeakyReLU(Module):
    def __init__(self, slope = 0.03):
        super(LeakyReLU, self).__init__()
            
        self.slope = slope
        
    def updateOutput(self, input):
        # Your code goes here. ################################################
        
        self.output = np.where(input <= 0, self.slope * input, input)
        
        return  self.output
    
    def updateGradInput(self, input, gradOutput):
        # Your code goes here. ################################################
        
        current_input = np.where(input > 0, 1, input)
        # print(current_input)
        current_input = np.where(current_input < 0, self.slope, current_input)
        # # print(current_input)
        self.gradInput = np.multiply(gradOutput, current_input)
        # print(current_input)
        # print(self.gradInput)
        
        return self.gradInput
    
    def __repr__(self):
        return "LeakyReLU"
    
    
class ELU(Module):
    def __init__(self, alpha = 1.0):
        super(ELU, self).__init__()
        
        self.alpha = alpha
        
    def updateOutput(self, input):
        # Your code goes here. ################################################
        self.output = np.where(input <= 0, self.alpha * (np.exp(input) - 1), input)
        # print(self.output.shape)
        return  self.output
    
    def updateGradInput(self, input, gradOutput):
        # Your code goes here. ################################################
        current_input = np.where(input > 0, 1, input)
        current_input = np.where(current_input <= 0, self.alpha * np.exp(current_input), current_input)
        self.gradInput = current_input * gradOutput
        # print(self.gradInput.shape)
        return self.gradInput
    
    def __repr__(self):
        return "ELU"
    
    
class SoftPlus(Module):
    def __init__(self):
        super(SoftPlus, self).__init__()
    
    def updateOutput(self, input):
        # Your code goes here. ################################################
        
        self.output = np.log(1 + np.exp(input))
        
        return  self.output
    
    def updateGradInput(self, input, gradOutput):
        # Your code goes here. ################################################
        
        self.gradInput = np.exp(input) / (1 + np.exp(input)) * gradOutput
        
        return self.gradInput
    
    def __repr__(self):
        return "SoftPlus"
    
    
class Criterion(object):
    def __init__ (self):
        self.output = None
        self.gradInput = None
        
    def forward(self, input, target):
        """
            Given an input and a target, compute the loss function 
            associated to the criterion and return the result.
            
            For consistency this function should not be overrided,
            all the code goes in `updateOutput`.
        """
        return self.updateOutput(input, target)

    def backward(self, input, target):
        """
            Given an input and a target, compute the gradients of the loss function
            associated to the criterion and return the result. 

            For consistency this function should not be overrided,
            all the code goes in `updateGradInput`.
        """
        return self.updateGradInput(input, target)
    
    def updateOutput(self, input, target):
        """
        Function to override.
        """
        return self.output

    def updateGradInput(self, input, target):
        """
        Function to override.
        """
        return self.gradInput   

    def __repr__(self):
        """
        Pretty printing. Should be overrided in every module if you want 
        to have readable description. 
        """
        return "Criterion"
    
    
    
class MSECriterion(Criterion):
    def __init__(self):
        super(MSECriterion, self).__init__()
        
    def updateOutput(self, input, target):   
        self.output = np.sum(np.power(input - target,2)) / input.shape[0]
        return self.output 
 
    def updateGradInput(self, input, target):
        self.gradInput  = (input - target) * 2 / input.shape[0]
        return self.gradInput

    def __repr__(self):
        return "MSECriterion"
    
    
class ClassNLLCriterionUnstable(Criterion):
    EPS = 1e-15
    def __init__(self):
        a = super(ClassNLLCriterionUnstable, self)
        super(ClassNLLCriterionUnstable, self).__init__()
        
    def updateOutput(self, input, target): 
        
        # Use this trick to avoid numerical errors
        input_clamp = np.clip(input, self.EPS, 1 - self.EPS)
        
        # Your code goes here. ################################################
        
        self.output = - np.sum(np.multiply(target, np.log(input_clamp)))/input.shape[0]
        # print(self.output)
        
        return self.output

    def updateGradInput(self, input, target):
        
        # Use this trick to avoid numerical errors
        input_clamp = np.clip(input, self.EPS, 1 - self.EPS)
                
        # Your code goes here. ################################################
        
        self.gradInput = - (target / input_clamp)/input.shape[0]
        # print(self.gradInput)
        
        return self.gradInput
    
    def __repr__(self):
        return "ClassNLLCriterionUnstable"
    
    
class ClassNLLCriterion(Criterion):
    def __init__(self):
        a = super(ClassNLLCriterion, self)
        super(ClassNLLCriterion, self).__init__()
        
    def updateOutput(self, input, target): 
        # Your code goes here. ################################################
        
        self.output = - np.sum(np.multiply(target, input))/input.shape[0]
        
        return self.output

    def updateGradInput(self, input, target):
        # Your code goes here. ################################################
        
        self.gradInput = - target/input.shape[0]
        
        return self.gradInput
    
    def __repr__(self):
        return "ClassNLLCriterion"
    
    
def sgd_momentum(variables, gradients, config, state):  
    # 'variables' and 'gradients' have complex structure, accumulated_grads will be stored in a simpler one
    state.setdefault('accumulated_grads', {})
    
    var_index = 0 
    for current_layer_vars, current_layer_grads in zip(variables, gradients): 
        for current_var, current_grad in zip(current_layer_vars, current_layer_grads):
            
            old_grad = state['accumulated_grads'].setdefault(var_index, np.zeros_like(current_grad))
            
            np.add(config['momentum'] * old_grad, config['learning_rate'] * current_grad, out=old_grad)
            
            current_var -= old_grad
            var_index += 1     
            
            
def adam_optimizer(variables, gradients, config, state):  
    # 'variables' and 'gradients' have complex structure, accumulated_grads will be stored in a simpler one
    state.setdefault('m', {})  # first moment vars
    state.setdefault('v', {})  # second moment vars
    state.setdefault('t', 0)   # timestamp
    state['t'] += 1
    for k in ['learning_rate', 'beta1', 'beta2', 'epsilon']:
        assert k in config, config.keys()
    
    var_index = 0 
    lr_t = config['learning_rate'] * np.sqrt(1 - config['beta2']**state['t']) / (1 - config['beta1']**state['t'])
    for current_layer_vars, current_layer_grads in zip(variables, gradients): 
        for current_var, current_grad in zip(current_layer_vars, current_layer_grads):
            var_first_moment = state['m'].setdefault(var_index, np.zeros_like(current_grad))
            var_second_moment = state['v'].setdefault(var_index, np.zeros_like(current_grad))
            
            # <YOUR CODE> #######################################
            # update `current_var_first_moment`, `var_second_moment` and `current_var` values
            #np.add(... , out=var_first_moment)
            #np.add(... , out=var_second_moment)
            #current_var -= ...
            
            np.add(config['beta1'] * var_first_moment, (1 - config['beta1']) * current_grad, out=var_first_moment)
            np.add(config['beta2'] * var_second_moment, (1 - config['beta2']) * np.multiply(current_grad, current_grad), out=var_second_moment)
            current_var -= lr_t * var_first_moment / (np.sqrt(var_second_moment) + config['epsilon'])
            # print(current_var)
            
            # small checks that you've updated the state; use np.add for rewriting np.arrays values
            assert var_first_moment is state['m'].get(var_index)
            assert var_second_moment is state['v'].get(var_index)
            var_index += 1

```

### Test 0: Initialization (0.01 points)

<!-- grade_id: cell-3473b7b6ffd64d07, points: 0.01, grade: true, solution: false -->
```python
# do not change this cell
import sys

import numpy
import numpy as np
import unittest

import collections
import pickle
import io


numpy.variance = numpy.var

```

### Test 1: Linear layer (0.04 points)

<!-- grade_id: cell-e3503c286039ec55, points: 0.04, grade: true, solution: false -->
```python
batch_size, n_in, n_out = 2, 3, 4
for d in data['test_Linear']:
    # layers initialization
    custom_layer = Linear(n_in, n_out)
    custom_layer.W = d['init_w']
    custom_layer.b = d['init_w_b']

    layer_input = d['input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    gt_output = d['output']
    assert np.allclose(gt_output, custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    gt_layer_grad = d['grad_output']
    assert np.allclose(gt_layer_grad, custom_layer_grad, atol=1e-6)

    # 3. check layer parameters grad
    custom_layer.accGradParameters(layer_input, next_layer_grad)
    weight_grad = custom_layer.gradW
    bias_grad = custom_layer.gradb
    torch_weight_grad = d['w_grad']
    torch_bias_grad = d['b_grad']
    assert np.allclose(torch_weight_grad, weight_grad, atol=1e-6)
    assert np.allclose(torch_bias_grad, bias_grad, atol=1e-6)
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_170361/3689530456.py in <module>
      1 batch_size, n_in, n_out = 2, 3, 4
----> 2 for d in data['test_Linear']:
      3     # layers initialization
      4     custom_layer = Linear(n_in, n_out)
      5     custom_layer.W = d['init_w']

NameError: name 'data' is not defined
```

### Test 2: Softmax (0.1 points)

<!-- grade_id: cell-e2c4124a6f815118, points: 0.1, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for d in data['test_SoftMax']:
    # layers initialization
    custom_layer = SoftMax()

    layer_input = d['input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    gt_output = d['output']
    assert np.allclose(gt_output, custom_layer_output, atol=1e-5)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    assert np.allclose(d['grad_output'], custom_layer_grad, atol=1e-5)
```

### Test 3: LogSoftMax (0.1 points)

<!-- grade_id: cell-69473387a23d8dff, points: 0.1, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for d in data['test_LogSoftMax']:
    # layers initialization
    custom_layer = LogSoftMax()

    layer_input = d['input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    gt_output = d['output']
    np.allclose(gt_output, custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    gt_grad = d['grad_output']
    np.allclose(gt_grad, custom_layer_grad, atol=1e-6)
```

### Test 4: BatchNormalization (0.1 points)

<!-- grade_id: cell-e5f59af66e6c111b, points: 0.1, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 32, 16
for d in data['test_BatchNormalization']:
    # layers initialization
    alpha = 0.9
    custom_layer = BatchNormalization(alpha)
    custom_layer.train()
    init_moving_mean = d['init_moving_mean']
    init_moving_variance = d['init_moving_variance']
    custom_layer.moving_mean = init_moving_mean
    custom_layer.moving_variance = init_moving_variance

    layer_input = d['input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    gt_output = d['output']
    np.allclose(gt_output, custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    gt_grad = d['grad_output']
    # please, don't increase `atol` parameter, it's garanteed that you can implement batch norm layer
    # with tolerance 1e-5
    np.allclose(gt_grad, custom_layer_grad, atol=1e-5)

    # 3. check moving mean
    gt_running_mean = d['gt_moving_mean']
    gt_running_var = d['gt_moving_var']
    np.allclose(custom_layer.moving_mean, gt_running_mean)
    # we don't check moving_variance because pytorch uses slightly different formula for it:
    # it computes moving average for unbiased variance (i.e var*N/(N-1))
    #self.assertTrue(np.allclose(custom_layer.moving_variance, torch_layer.running_var.numpy()))

    # 4. check evaluation mode
    custom_layer.moving_variance = gt_running_var
    custom_layer.evaluate()
    custom_layer_output = custom_layer.updateOutput(layer_input)
    eval_output = d['eval_output']
    np.allclose(eval_output, custom_layer_output, atol=1e-6)
```

### Test 5: Sequential (0.1 points)

<!-- grade_id: cell-1a9c1e3609ed9aab, points: 0.1, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for d in data['test_Sequential']:
    # layers initialization
    alpha = 0.9
    custom_layer = Sequential()
    bn_layer = BatchNormalization(alpha)
    bn_layer.moving_mean = d['init_moving_mean']
    bn_layer.moving_variance = d['init_moving_var']

    custom_layer.add(bn_layer)
    scaling_layer = ChannelwiseScaling(n_in)
    scaling_layer.gamma = d['init_gamma']
    scaling_layer.beta = d['init_beta']
    custom_layer.add(scaling_layer)
    custom_layer.train()

    layer_input = d['input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    assert np.allclose(d['gt_output'], custom_layer_output, atol=1e-5)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.backward(layer_input, next_layer_grad)
    assert np.allclose(d['gt_grad'], custom_layer_grad, atol=1e-5)

    # 3. check layer parameters grad
    weight_grad, bias_grad = custom_layer.getGradParameters()[1]
    torch_weight_grad = d['gt_weight_grad']
    torch_bias_grad = d['gt_bias_grad']
    assert np.allclose(torch_weight_grad, weight_grad, atol=1e-5)
    assert np.allclose(torch_bias_grad, bias_grad, atol=1e-5)
```

### Test 6: Dropout (0.075 points)

<!-- grade_id: cell-1ddb0377b6c68deb, points: 0.15, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for _ in range(100):
    # layers initialization
    p = np.random.uniform(0.3, 0.7)
    layer = Dropout(p)
    layer.train()

    layer_input = np.random.uniform(-5, 5, (batch_size, n_in)).astype(np.float32)
    next_layer_grad = np.random.uniform(-5, 5, (batch_size, n_in)).astype(np.float32)

    # 1. check layer output
    layer_output = layer.updateOutput(layer_input)
    assert np.all(np.logical_or(np.isclose(layer_output, 0),
                                np.isclose(layer_output*(1.-p), layer_input)))

    # 2. check layer input grad
    layer_grad = layer.updateGradInput(layer_input, next_layer_grad)
    assert np.all(np.logical_or(np.isclose(layer_grad, 0),
                                np.isclose(layer_grad*(1.-p), next_layer_grad)))

    # 3. check evaluation mode
    layer.evaluate()
    layer_output = layer.updateOutput(layer_input)
    assert np.allclose(layer_output, layer_input)

    # 4. check mask
    p = 0.0
    layer = Dropout(p)
    layer.train()
    layer_output = layer.updateOutput(layer_input)
    assert np.allclose(layer_output, layer_input)

    p = 0.5
    layer = Dropout(p)
    layer.train()
    layer_input = np.random.uniform(5, 10, (batch_size, n_in)).astype(np.float32)
    next_layer_grad = np.random.uniform(5, 10, (batch_size, n_in)).astype(np.float32)
    layer_output = layer.updateOutput(layer_input)
    zeroed_elem_mask = np.isclose(layer_output, 0)
    layer_grad = layer.updateGradInput(layer_input, next_layer_grad)
    assert np.all(zeroed_elem_mask == np.isclose(layer_grad, 0))

    # 5. dropout mask should be generated independently for every input matrix element, not for row/column
    batch_size, n_in = 1000, 1
    p = 0.8
    layer = Dropout(p)
    layer.train()

    layer_input = np.random.uniform(5, 10, (batch_size, n_in)).astype(np.float32)
    layer_output = layer.updateOutput(layer_input)
    assert np.sum(np.isclose(layer_output, 0)) != layer_input.size

    layer_input = layer_input.T
    layer_output = layer.updateOutput(layer_input)
    assert np.sum(np.isclose(layer_output, 0)) != layer_input.size
```

### Test 7: LeakyReLU (0.05 points)

<!-- grade_id: cell-da6c6ccf757a0621, points: 0.05, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for d in data['test_LeakyReLU']:
    # layers initialization
    slope = d['slope']
    custom_layer = LeakyReLU(slope)

    layer_input = d['layer_input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    assert np.allclose(d['gt_output'], custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    assert np.allclose(d['gt_grad'], custom_layer_grad, atol=1e-6)
```

### Test 8: ELU (0.075 points)

<!-- grade_id: cell-301cefbceb12834e, points: 0.075, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for d in data['test_ELU']:
    # layers initialization
    alpha = 1.0
    custom_layer = ELU(alpha)

    layer_input = d['layer_input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    assert np.allclose(d['gt_output'], custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    assert np.allclose(d['gt_grad'], custom_layer_grad, atol=1e-6)
```

### Test 9: SoftPlus (0.1 points)

<!-- grade_id: cell-f804d9bf6f70e562, points: 0.1, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for d in data['test_SoftPlus']:
    # layers initialization
    custom_layer = SoftPlus()

    layer_input = d['layer_input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    assert np.allclose(d['gt_output'], custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    assert np.allclose(d['gt_grad'], custom_layer_grad, atol=1e-6)

```

### Test 10: ClassNLLCriterionUnstable (0.1 points)

<!-- grade_id: cell-76b1d3d075e9752b, points: 0.1, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for d in data['test_ClassNLLCriterionUnstable']:
    # layers initialization
    custom_layer = ClassNLLCriterionUnstable()

    layer_input = d['layer_input']
    target = d['target']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input, target)
    assert np.allclose(d['gt_output'], custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, target)
    assert np.allclose(d['gt_grad'], custom_layer_grad, atol=1e-6)
```

### Test 11: ClassNLLCriterion (0.05 points)

<!-- grade_id: cell-cbc87cb79435ffda, points: 0.05, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 4
for d in data['test_ClassNLLCriterion']:
    # layers initialization
    custom_layer = ClassNLLCriterion()

    layer_input = d['layer_input']
    target = d['target']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input, target)
    assert np.allclose(d['gt_output'], custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, target)
    assert np.allclose(d['gt_grad'], custom_layer_grad, atol=1e-6)
```

### Test 12: Adam (0.1 points)

<!-- grade_id: cell-31b7bd19c0cfd8f5, points: 0.1, grade: true, solution: false -->
```python
state = {}
config = {'learning_rate': 1e-3, 'beta1': 0.9, 'beta2':0.999, 'epsilon':1e-8}
variables = [[np.arange(10).astype(np.float64)]]
gradients = [[np.arange(10).astype(np.float64)]]
adam_optimizer(variables, gradients, config, state)
assert np.allclose(state['m'][0], np.array([0. , 0.1, 0.2, 0.3, 0.4, 0.5,
                                                     0.6, 0.7, 0.8, 0.9]))
assert np.allclose(state['v'][0], np.array([0., 0.001, 0.004, 0.009, 0.016, 0.025,
                                                     0.036, 0.049, 0.064, 0.081]))
assert state['t'] == 1
assert np.allclose(variables[0][0], np.array([0., 0.999, 1.999, 2.999, 3.999, 4.999,
                                                       5.999, 6.999, 7.999, 8.999]))
adam_optimizer(variables, gradients, config, state)
assert np.allclose(state['m'][0], np.array([0., 0.19, 0.38, 0.57, 0.76, 0.95, 1.14,
                                                     1.33, 1.52, 1.71]))
assert np.allclose(state['v'][0], np.array([0., 0.001999, 0.007996, 0.017991,
                                                     0.031984, 0.049975, 0.071964, 0.097951,
                                                     0.127936, 0.161919]))
assert state['t'] == 2
assert np.allclose(variables[0][0], np.array([0., 0.998, 1.998, 2.998, 3.998, 4.998,
                                                       5.998, 6.998, 7.998, 8.998]))
```

### Test 13: Conv2d (0.35 points)

<!-- grade_id: cell-b63afa10ee124617, points: 0.35, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in, n_out = 2, 3, 4
h,w = 5,6
kern_size = 3
for d in data['test_Conv2d']:
    # layers initialization
    custom_layer = Conv2d(n_in, n_out, kern_size)
    custom_layer.W = d['init_w']
    custom_layer.b = d['init_b']

    layer_input = d['input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    assert np.allclose(d['gt_output'], custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    assert np.allclose(d['gt_grad'], custom_layer_grad, atol=1e-6)

    # 3. check layer parameters grad
    custom_layer.accGradParameters(layer_input, next_layer_grad)
    weight_grad = custom_layer.gradW
    bias_grad = custom_layer.gradb
    assert np.allclose(d['gt_weight_grad'], weight_grad, atol=1e-6, )
    assert np.allclose(d['gt_bias_grad'], bias_grad, atol=1e-6)

```

### Test 13: MaxPool2d (0.15 points)

<!-- grade_id: cell-1392d1511ad4da77, points: 0.15, grade: true, solution: false -->
```python
np.random.seed(42)

batch_size, n_in = 2, 3
h,w = 4,6
kern_size = 2
for d in data['test_MaxPool2d']:
    # layers initialization
    custom_layer = MaxPool2d(kern_size)

    layer_input = d['input']
    next_layer_grad = d['next_layer_grad']

    # 1. check layer output
    custom_layer_output = custom_layer.updateOutput(layer_input)
    assert np.allclose(d['gt_output'], custom_layer_output, atol=1e-6)

    # 2. check layer input grad
    custom_layer_grad = custom_layer.updateGradInput(layer_input, next_layer_grad)
    assert np.allclose(d['gt_grad'], custom_layer_grad, atol=1e-6)
```